Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jianzhang Chen"
COLLABORATORS = ""
```

---

# Datacleaning with Python
The goal of this assignment is to get you started with Pandas. We will use the same dataset as R assignment and learn to do following things in Pandas:
1. Read a csv file.
2. Column wise and row wise transformations with `apply` method.
3. Analyse summary of data with `describe` and `groupby`.



```python
import numpy as np
import pandas as pd
from datetime import datetime
```

## Loading data.

```python
# If all values in a column are null, Pandas loads it as float column. 
# We specifically tell Pandas to load following columns as datatype object.
# object (roughly) coressponds to string in Pandas.
menu_dtypes = {"keywords" : np.object, "language" : np.object, "location_type" : np.object}
# Read csv.
menu = pd.read_csv("~/data_readonly/A5/Menu.csv", dtype=menu_dtypes)
```

```python
# Here is a preview of the dataset you just loaded. Scroll right to take a look at the date column.
menu
```

Output:
```
id                      name  \
0      12463                       NaN   
1      12464                       NaN   
2      12465                       NaN   
3      12466                       NaN   
4      12467                       NaN   
5      12468                       NaN   
6      12469                       NaN   
7      12470                       NaN   
8      12471                       NaN   
9      12472                       NaN   
10     12473                       NaN   
11     12474                       NaN   
12     12475                       NaN   
13     12476                       NaN   
14     12477                       NaN   
15     12478                       NaN   
16     12479                       NaN   
17     12480                       NaN   
18     12481                       NaN   
19     12482                       NaN   
20     12483                       NaN   
21     12484                       NaN   
22     12485                       NaN   
23
```

Output of next cell will tell you datatype of different columns in the dataframe. As stated earlier, object means string.  
As you will in the output of this cell, date coulmn's datatype is string intead of datetime. This is because, some of the dates are ill formed e.g. some dates have year 1 or 190 which are obviously incorrect. Next we will identify such dates, remove them and will add a new date column with correct datatype.

```python
menu.dtypes
```

Output:
```
id                       int64
name                    object
sponsor                 object
event                   object
venue                   object
place                   object
physical_description    object
occasion                object
notes                   object
call_number             object
keywords                object
language                object
date                    object
location                object
location_type           object
currency                object
currency_symbol         object
status                  object
page_count               int64
dish_count               int64
dtype: object
```

### Extracting year, month and day from date.
We will write three functions to extract year, month and day from date respectively and then apply these functions on each row.

<!-- grade_id: cell-f08284e5f2e32d2c, points: 0, grade: false, solution: true -->
```python
def extract_year(date_str):
    '''
    Given a string of format of "yyyy-MM-dd" extract year. Use split()
    If the input in null, return null. Code is already provided for this part. You just have to modify else part
    
    Don't forget to convert to integer using int() before returning.
    '''
    if pd.isnull(date_str):
        return np.NaN
    else:
        # Below line just returns 0 always. You have to change it to extract year from date_str using split.
        # Make sure to change to int type before returning. Tests will not pass otherwise.
        # Your implementation should not assume that year is always 4 digits. Some ill-formed dates could have less/more.
        year = date_str.split('-')[0]
        return int(year)


```

<!-- grade_id: cell-f17153f303ff0176, points: 2, grade: true, solution: false -->
```python
# [2 points] Testing if above function is implemented correctly.
assert extract_year("1990-01-31") == 1990
assert type(extract_year("1990-01-31")) == int
assert extract_year("05-01-31") == 5
assert type(extract_year("05-01-31")) == int
```

<!-- grade_id: cell-71c4a773f6cc752f, points: 8, grade: true, solution: false -->
```python
# [8 points] A test very similar to previous but will only run after submission.

```

<!-- grade_id: cell-5d21638bfacee0e5, points: 0, grade: false, solution: true -->
```python
def extract_month(date_str):
    '''
    Given a string of format of "yyyy-MM-dd" extract month. Use split()
    If the input in null, return null. Code is already provided for this part. You just have to modify else part
    
    Don't forget to convert to integer using int() before returning.
    '''
    if pd.isnull(date_str):
        return np.NaN
    else:
        # Below line just returns 0 always. You have to change it to extract month from date_str using split.
        # Make sure to change to int type before returning. Tests will not pass otherwise.
        # Your implementation should not assume that month is always 2 digits. Some ill-formed dates could have less/more.
        return int(date_str.split('-')[1])



```

<!-- grade_id: cell-7219457cbd4c2daf, points: 2, grade: true, solution: false -->
```python
# [2 points] Testing if above function is implemented correctly.
assert extract_month("1990-01-31") == 1
assert type(extract_month("1990-01-31")) == int
assert extract_month("05-7-31") == 7
assert type(extract_month("05-7-31")) == int
```

<!-- grade_id: cell-d0d1b8485f1a711a, points: 8, grade: true, solution: false -->
```python
# [8 points] A test very similar to previous but will only run after submission.

```

<!-- grade_id: cell-2a0b4e5d06e9f67a, points: 0, grade: false, solution: true -->
```python
def extract_day(date_str):
    '''
    Given a string of format of "yyyy-MM-dd" extract day. Use split()
    If the input in null, return null. Code is already provided for this part. You just have to modify else part
    
    Don't forget to convert to integer using int() before returning.
    '''
    if pd.isnull(date_str):
        return np.NaN
    else:
        # Below line just returns 0 always. You have to change it to extract day from date_str using split.
        # Make sure to change to int type before returning. Tests will not pass otherwise.
        # Your implementation should not assume that day is always 2 digits. Some ill-formed dates could have less/more.
        return int(date_str.split('-')[2])



```

<!-- grade_id: cell-c7c79374846fa844, points: 2, grade: true, solution: false -->
```python
# [2 points] Testing if above function is implemented correctly.
assert extract_day("1990-01-31") == 31
assert type(extract_day("1990-01-31")) == int
assert extract_day("1990-7-3") == 3
assert type(extract_day("1990-7-3")) == int
```

<!-- grade_id: cell-f281c35f238187cb, points: 8, grade: true, solution: false -->
```python
# [8 points] A test very similar to previous but will only run after submission.

```

### Apply above functions to each date.
`apply` methods are the recommended to apply some tranformation single column or even group of columns too. You should almost never use loops while working with dataframes. 

```python
# Add three new columns to the dataframe.
menu["year"] = menu["date"].apply(extract_year)
menu["month"] = menu["date"].apply(extract_month)
menu["day"] = menu["date"].apply(extract_day)
```

<!-- grade_id: cell-f74227bbdad912f2, points: 0, grade: false, solution: true -->
```python
# describe method on dataframe is used to look at some stats like min, max, mean, median etc.
# call describe method on menu on next line and answer next few questions.

```

<!-- grade_id: cell-e2e8d8b0140816f6, points: 0, grade: false, solution: true -->
```python
# Uncomment following line and write the max value of year as reported by describe method.
# You don't have to write 0s after decimal.
max_value_of_year = int(max(menu['year']))
max_value_of_year
```

Output:
```
2928
```

<!-- grade_id: cell-10dfcaea7a6fb5be, points: 5, grade: true, solution: false -->
```python
# [5 points] This test will check max_value_of_year and will run after submission.

```

<!-- grade_id: cell-ecc987808c5234dd, points: 0, grade: false, solution: true -->
```python
# Uncomment following line and write the min value of year as reported by describe method.
# You don't have to write 0s after decimal.
min_value_of_year = int(min(menu['year']))
min_value_of_year

```

Output:
```
1
```

<!-- grade_id: cell-5794e69e22574042, points: 5, grade: true, solution: false -->
```python
# [5 points] This test will check min_value_of_year and will run after submission.

```

<!-- grade_id: cell-699bb1f86ff93fa1, points: 0, grade: false, solution: true -->
```python
# Uncomment following line and write the 75th percentile value of day as reported by describe method.
# You don't have to write 0s after decimal.
percentile_75_day = int(menu['day'].quantile(.75))
percentile_75_day

```

Output:
```
22
```

<!-- grade_id: cell-be510ccbf9cf3916, points: 5, grade: true, solution: false -->
```python
# [5 points] This test will check percentile_75_day and will run after submission.

```

<!-- grade_id: cell-274b4b9f726afe5e, points: 0, grade: false, solution: true -->
```python
# Uncomment following line and write the 25th percentile value of year as reported by describe method.
# You don't have to write 0s after decimal.
percentile_25_year = int(menu['year'].quantile(.25))
percentile_25_year


```

Output:
```
1900
```

<!-- grade_id: cell-eb77ebd10cd96481, points: 5, grade: true, solution: false -->
```python
# [5 points] This test will check percentile_25_year and will run after submission.

```

### Next we will identify the dates which are most probably wrong e.g. dates older than 1800. We will replace all such dates with null values.

Let's start with identifying number of rows for each year. Use Pandas `groupby` method for this question and create a dictionary where key is year and value is number of rows in which that year appears e.g. if number of rows with year 1990 are 2 and number of rows with year 2000 are 5, your dictionary should be like the example shown in next cell. However, you will not hardcode like the example but compute it using `groupby`. 

Hint: It should take you only line of code if you do in proper Pandas way.  
Hint: Aggregation functions on `groupby` aggregate for each column individually and only considers rows where that column in non-null. Therefore, after grouping and before counting number of rows, you will have to filter for a column which is always non-null like `id` column. This hint is intentionally vague so you spend sometime with group by and understand exactly how it works by trying out different things.

```python
example_dictionary = {1990.0 : 2, 2000.0 : 5}
example_dictionary
```

Output:
```
{1990.0: 2, 2000.0: 5}
```

<!-- grade_id: cell-a4d0266466a47218, points: 0, grade: false, solution: true -->
```python
# Modify the following line and set the variable rows_per_year to a dictionary.
rows_per_year = dict(menu.groupby('year')['id'].count())
rows_per_year

```

Output:
```
{1.0: 2,
 190.0: 1,
 1091.0: 1,
 1851.0: 3,
 1852.0: 2,
 1853.0: 1,
 1854.0: 3,
 1855.0: 1,
 1856.0: 5,
 1857.0: 1,
 1858.0: 3,
 1859.0: 3,
 1861.0: 1,
 1862.0: 3,
 1864.0: 3,
 1865.0: 7,
 1866.0: 6,
 1869.0: 1,
 1873.0: 1,
 1880.0: 20,
 1881.0: 33,
 1882.0: 25,
 1883.0: 14,
 1884.0: 19,
 1885.0: 21,
 1886.0: 19,
 1887.0: 29,
 1888.0: 46,
 1889.0: 57,
 1890.0: 57,
 1891.0: 95,
 1892.0: 99,
 1893.0: 74,
 1894.0: 74,
 1895.0: 111,
 1896.0: 125,
 1897.0: 169,
 1898.0: 210,
 1899.0: 468,
 1900.0: 3326,
 1901.0: 2263,
 1904.0: 1,
 1905.0: 336,
 1906.0: 571,
 1907.0: 504,
 1908.0: 78,
 1909.0: 3,
 1910.0: 280,
 1912.0: 264,
 1913.0: 104,
 1914.0: 804,
 1915.0: 269,
 1916.0: 173,
 1917.0: 592,
 1918.0: 396,
 1919.0: 131,
 1920.0: 122,
 1921.0: 123,
 1922.0: 1,
 1924.0: 14,
 1925.0: 30,
 1926.0: 30,
 1927.0: 38,
 1928.0: 50,
 1929.0: 37,
 1930.0: 5,
 1931.0: 27,
 1932.0: 54,
 1933.0: 393,
 1934.0: 39,
 1935.0: 77,
 1936.0: 1,
 1937.0: 209,
 1938.0: 169,
 1939.0: 206,
 1940.0: 125,
 1941.0: 102
```

<!-- grade_id: cell-11ec2648226f39b1, points: 30, grade: true, solution: false -->
```python
# [30 points] This will test if the dictionary is correct and will only run after submission.

```

Examine dictionary for outliers i.e. if there are rows with year smaller than `1800` or greater than `2018`. This exercise will convince you that there are some rows with improbable dates, and these dates are most probably data entering errors.  
Do the same thing for `month` and `day` columns, and observe if there are any rows with potentially wrong day or month i.e. if there is even a single row with day less than 1 or greater than 31, and if there is even a single row with month less than 1 or greater than 12. Convince yourself that there are no rows with month or day outside legal range. 

Examining year, month and day will only tell you if each of them are correct individually. However, this doesn't mean that complete date is correct e.g. `2018-02-30` is not a valid date but all individual components are correct in their own right. Next, we will write some code to check if complete dates are correct.

```python
# Let's replace year, month and day with null whenever year is less than 1800 or greater than 2018.
# You don't have to write any code for this, we are providing it. Make sure you understand this code.
menu.loc[menu["year"] < 1800, "year"] = np.nan
menu.loc[menu["year"] > 2018, "year"] = np.nan

menu.loc[menu["year"] < 1800, "day"] = np.nan
menu.loc[menu["year"] > 2018, "day"] = np.nan

menu.loc[menu["year"] < 1800, "month"] = np.nan
menu.loc[menu["year"] > 2018, "month"] = np.nan
```

We will use three columns, `day`, `year` and `month` to create a new date column. If some `day`, `month`, `year` triple is not compatible, we will not be able to create a date object out of them and hence we will find "unclean"  data.  
Just like before we will write a function to create date given a triple, and call this method for each row using `apply`

<!-- grade_id: cell-60119476d7d201c4, points: 0, grade: false, solution: true -->
```python
# Complete the function below. Make sure to return datetime object from this function.
def create_date(year, month, day):
    if pd.isnull(year) or pd.isnull(month) or pd.isnull(day):
        return np.nan
    else:
        year = int(year)
        month = int(month)
        day = int(day)
        # Change next line to create datetime from year, month and day.
        return datetime(year=year,month=month,day=day) # Change this line.
    
    

```

<!-- grade_id: cell-13891c65c61d8a6c, points: 20, grade: true, solution: false -->
```python
# [20 points] This test will check if previous function is implemented correctly. This will run only after submission.

```

```python
# We are going to use the function you just wrote to create a new column.
# If the following line does not fail, then it means that all triples were valid dates otherwise there was some error.
menu["date_cleaned"] = menu.apply(lambda row: create_date(row["year"], row["month"], row["day"]), axis=1)
```

```python
# Check that new column has correct datatype.
menu.dtypes
```

Output:
```
id                               int64
name                            object
sponsor                         object
event                           object
venue                           object
place                           object
physical_description            object
occasion                        object
notes                           object
call_number                     object
keywords                        object
language                        object
date                            object
location                        object
location_type                   object
currency                        object
currency_symbol                 object
status                          object
page_count                       int64
dish_count                       int64
year                           float64
month                          float64
day                            float64
date_cleaned            datetime64[ns]
dtype: object
```

```python

```